![D590](https://homes.luddy.indiana.edu/classes/spring2025/dsci/d590-skavousi/files/banner.png)
# Assignment 02: Conditionals and Loops

```python
# This cell tells Jupyter to provide detailed debugging information
# when a runtime error occurs. Run it before working on the questions.

%xmode Verbose
```

# Returning a Tuple Example
The function below shows how to return a letter grade and level as a **tuple**. While we'll cover tuples in more detail later in the course, for now you just need to return values in this format for the first question for the autograder to work correctly.

```python
def grade_score(score):
    grade = None
    level = None
    
    if score >= 90:
        grade = 'A'
        level = "Excellent"
    elif score >= 80:
        grade = 'B'
        level = "Above Average"
    elif score >= 70:
        grade = 'C'
        level = "Average"
    elif score >= 60:
        grade = 'D'
        level = "Below Average"
    else:
        grade = 'F'
        level = "Failing"

    # Here we are returning a tuple
    # The parenthesis are required for a tuple
    return (grade, level) 
```

Here's an example usage of `grade_score`. We can see it returns a **tuple** of `('D', 'Below Average')`

```python
result = grade_score(63)
print(result)
print(type(result))
```

Here's how we'll autograde your functions when we ask you to return a tuple.

```python
assert grade_score(74) == ('C', 'Average')
```

<div class="alert alert-block alert-info">

## Question 1 (50 points):

### Overview
The body mass index (BMI) is a statistic developed by Adolphe Quetelet in the 1900’s for evaluating body mass. It uses the same formula for men as for women and children.
According to http://www.bmi3d.com/formula.html the body mass index is calculated as kilograms divided by height in meters squared:
$$
\text{BMI}=\frac{x}{y^2}
$$

where $x$ is body weight in kg and $y$ is height in meters.

A BMI in the range 18.5 to 24.9 (inclusive) is considered normal.

### Implementation
Write a function `body_mass_index(height, weight)` that:
- Takes height in inches and weight in pounds
- Converts height to meters, and weight to kilograms
- Returns a tuple `(bmi, message)`
- BMI should be rounded to one decimal place
- Message should be 'underweight', 'normal', or 'overweight' based on BMI range
- Returns `(-1, "error message")` if height or weight is negative

### Important
- For consistency in test cases, please use exactly `0.0254` for inches-to-meters and `0.454` for pounds-to-kilograms conversions.

### Hints
- Our autograder tests rely on exact text matching. Make sure you closely follow the implementation instructions

### Scoring
- Correct mathematical implementation (conversion, rounding) [20]
- Correct function implementation (parameters, return statement) [20]
- Proper comments, spacing, formatting, and follosing naming conventions for variables [10]
</div>


<!-- grade_id: question_1_student, points: 0, grade: false, solution: true -->
```python
def body_mass_index(height, weight):
    """Calculate Body Mass Index (BMI) from height (inches) and weight (pounds).
    
    Returms a tuple: (bmi, message)
    - bmi: float rounded to 1 decimal 
    - message: "underweight', 'normal', 'overweight'
    - Returns (-1, "error messahe") if invalid input
    """
    #Error handling
    if height <= 0 or weight <= 0:
        return (-1, "error message")
        
    #Convert units
    height_m = height * 0.0254
    weight_kg = weight * 0.454
    
    #Calculate BMI
    bmi = weight_kg/ (height_m **2)
    bmi = round(bmi, 1)
    
    #Determine category
    if bmi < 18.5:
        message = "underweight"
    elif 18.5 <= bmi <= 24.9: 
        message = "normal"
    else: 
        message = "overweight"

    return (bmi, message)
        
```

```python
# You can run this cell after writing your solutions to
# check it against some basic test cases. Do not delete
# or modify the contents of this cell.
assert body_mass_index(-70, 200)[0] == -1
assert body_mass_index(70, -200)[0] == -1
assert body_mass_index(70, 100) == (14.4, 'underweight')
assert body_mass_index(70, 200) == (28.7, 'overweight')
assert body_mass_index(70, 128) == (18.4, 'underweight')
assert body_mass_index(70, 174) == (25.0, 'overweight')

print("All tests passed!")
```

Output:
```
All tests passed!
```

<!-- grade_id: body_mass_index_return, points: 15, grade: true, solution: false -->
```python
# This cell is intentionally left blank, do NOT delete this cell or modify it.
```

<!-- grade_id: body_mass_index_math_threshold, points: 25, grade: true, solution: false -->
```python
# This cell is intentionally left blank, do NOT delete this cell or modify it.
```

<!-- grade_id: body_mass_index_math, points: 0, grade: true, solution: false -->
```python
# This cell is intentionally left blank, do NOT delete this cell or modify it.
```

<!-- grade_id: body_mass_index_math_function, points: 10, grade: true, solution: false -->
```python
# This cell is intentionally left blank, do NOT delete this cell or modify it.
```

<div class="alert alert-block alert-info">

## Question 2 (50 points):

Taken directly from "Think Python, 2nd Edition. https://greenteapress.com/thinkpython2/html/thinkpython2008.html#squareroot" with minor editing. 

### Overview
To understand what an algorithm is, it might help to start with something that is not an algorithm. When you learned to multiply single-digit numbers, you probably memorized the multiplication table. In effect, you memorized 100 specific solutions. That kind of knowledge is not algorithmic.

But if you were “lazy”, you might have learned a few tricks. For example, to find the product of $n$ and 9, you can write $n−1$ as the first digit and $10−n$ as the second digit. This trick is a general solution for multiplying any single-digit number by 9. That’s an algorithm!

Similarly, the techniques you learned for addition with carrying, subtraction with borrowing, and long division are all algorithms. One of the characteristics of algorithms is that they do not require any intelligence to carry out. They are mechanical processes where each step follows from the last according to a simple set of rules.

Executing algorithms is boring, but designing them is interesting, intellectually challenging, and a central part of computer science. 

Some of the things that people do naturally, without difficulty or conscious thought, are the hardest to express algorithmically. Understanding natural language is a good example. We all do it, but so far no one has been able to explain how we do it, at least not in the form of an algorithm.

Newton’s method is an example of an algorithm: it is a mechanical process for solving a category of problems (in this case, computing square roots).

### Implementation

Loops are often used in programs that compute numerical results by starting with an approximate answer and iteratively improving it.

For example, one way of computing square roots is Newton’s method. Suppose that you want to know the square root of $a$. If you start with almost any estimate, $x$, you can compute a better estimate with the following formula:
$$
y = \frac{x + \frac{a}{x}}{2}
$$

For example, if $a$ is 4 and x is 3:
```python
>>> a = 4
>>> x = 3
>>> y = (x + a/x) / 2
>>> y
2.16666666667
```
The result is closer to the correct answer ($\sqrt{4}$ = 2). If we repeat the process with the new estimate, it gets even closer:
```python
>>> x = y
>>> y = (x + a/x) / 2
>>> y
2.00641025641
```
After a few more updates, the estimate gets very very close:
```python
>>> x = y
>>> y = (x + a/x) / 2
>>> y
2.00001024003
>>> x = y
>>> y = (x + a/x) / 2
>>> y
2.00000000003
```
In general we don’t know ahead of time how many steps it takes to get to the right answer, but we know when we get there because the estimate stops changing:
```python
>>> x = y
>>> y = (x + a/x) / 2
>>> y
2.0
>>> x = y
>>> y = (x + a/x) / 2
>>> y
2.0
```
Rather than checking whether $x$ and $y$ are exactly equal, it is safer to use the built-in function `abs` to compute the absolute value, or magnitude, of the difference between them:
```python
if abs(y-x) < epsilon:
    ...
```

Define epsilon = 0.1

### Task
1. Write a function `mysqrt(a)` that takes `a` as a parameter, chooses a reasonable value of `x`, and returns an estimate of the square root of `a` using Newton's method. **Important:** Use `epsilon = 0.1`
2. Write a function `test_square_root(a)` that takes `a` as a parameter and returns a tuple `(x, y, z)` where:
   - `x` is the square root of `a` computed with `mysqrt`
   - `y` is the square root of `a` computed by `math.sqrt`
   - `z` is the absolute value of the difference between the two estimates `x` and `y`

### Important

Make sure the starting value of `x` is a reasonable value for each input. 
Using fixed values for `x` or unreasonable estimates will result in loss of points and also unnecessary computations.

### Scoring
- Correct `mysqrt(a)` math and function implementation with reasonable stating `x` value [25]
- Correct `test_square_root` implementation [15]
- Proper comments, spacing, formatting, and following naming conventions for variables [10]

</div>

<!-- grade_id: question_2_student, points: 0, grade: false, solution: true -->
```python
import math
def mysqrt(a):
    """ 
    Approximate the square root of a number using Newton's method. 
    The iteration stops when the change is less than epsilon = 0.1. 
    """
    if a <0: 
        return -1

    epsilon = 0.1
    x= a/2 if a >1 else 1

    while True:
        y = (x+a/x) /2
        if abs(y-x) <epsilon:
            return y
        x = y
        
def test_square_root(a):
    ''' 
    Compare mysqrt(a) with math.sqrt (a).

    Returns a tuple (x,y,z):
    -x: square root from mysqrt
    -y: square root from math.sqrt
    -z: absolute differemce 
    '''
    x= mysqrt (a)
    y = math.sqrt(a)
    z= abs(x-y)
    return (x,y,z)

    
```

<!-- grade_id: test_square_root_sc1, points: 20, grade: true, solution: false -->
```python
# You can run this cell after writing your solutions to
# check it against some basic test cases. Do not delete
# or modify the contents of this cell.
THRESHOLD = 0.01
assert abs(mysqrt(4) - 2) <= THRESHOLD
assert abs(mysqrt(16) - 4) <= THRESHOLD
assert abs(mysqrt(10) - 3.162277660168379) <= THRESHOLD
print("All tests passed!")
```

Output:
```
All tests passed!
```

<!-- grade_id: test_square_root, points: 30, grade: true, solution: false -->
```python
# This cell is intentionally left blank, do NOT delete this cell or modify it.
```

<div class="alert alert-block alert-success">  

## Instructions:

- Make sure your code is properly formatted, structured and commented. **The statements in the program have to be properly indented.**

- Do your best to follow the programming style suggested by PEP8, [tutorial here](https://realpython.com/python-pep8/). Follow the naming conventions for variables and functions.

- Make sure that your program compiles and runs successfully. A program that throws an error is usually scored 0 points.  
</div>